---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
free_preview: true
---

## Example coding exercise

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Insert exercise title here

```yaml
type: NormalExercise
key: d920164d58
xp: 100
```

Cree una función que realice la suma de un vectore de datos. 		

`@instructions`
Realizarlo por medio de la función function

`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}
###########
funcion<-function{
  
}
```

`@solution`
```{r}
funcion<-function(x){
  sum(x)
}
```

`@sct`
```{r}

```
